Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and University of Bristol ID:

```python
NAME = "Rafaela Ioannou"
UoB_ID = "um18532"
```

# Homework 6 

This homework continues our study of  number theoretic functions and notions that are relevant to developing cyrptographic algorithms.  

## Question 1: $m$-ary Conversion.

This question is a generalisation of the Question 1 of the tutorial. 

**(a) (1 mark)** Write a function `convert` that, given a pair of integers `(number,base)` as inputs, with `number` being non-negative and `base` being a member of the set $\{2,\dots,32\}$, returns a string representing `number` in `base`-ary notation. 
Note that 2-ary is just binary and that, for example, the 5-ary string representation of `43` is `'133'` since $43 = 1 \cdot 5^2 + 3 \cdot 5^1 + 3 \cdot 5^0$. 
 To represent the numbers $10,11,\dots,31$ in your string representation you should use the characters `'a','b',...,'u','v'`. (So, for example, `convert(59,12)` should return `'4b'`.) 

<!-- grade_id: cell-85ec2634b6599a0c, points: 0, grade: false, solution: true -->
```python
def convert(number,base):
    # YOUR CODE HERE
    #raise NotImplementedError()
    wt_message = 'Number input of wrong type'
    ng_message = 'Number input is negative'
    wt_base_message = 'Base imput of wrong type'
    small_base_message = 'Base imput is less than 2'
    large_base_message = 'Base imput is bigger than 32'
    if not type(number) is int: 
        print(wt_message)
        return wt_message
    if number < 0: 
        print(ng_message)
        return ng_message
    if not type(base) is int:
        print(wt_base_message)
        return wt_base_message
    if base < 2:
        print(small_base_message)
        return small_base_message
    if base > 32:
        print(large_base_message)
        return large_base_message
    #We have dealt with bad input
    #Now for the main part
    
    m_string = ''
    if number == 0:
        m_string = '0'
    while number > 0:
        if number%base == 0:
            m_string = '0' + m_string
        elif number%base < 10:
            N = number%base
            m_string = str(N) + m_string
        elif number%base >= 10:
            M = number%base
            M1 = chr(M+87)
            m_string = str(M1) + m_string
        number = number//base
    return m_string  


```

```python
# Change the value of base and max and test your functions with these values
# in the next cell 
base = 28
max = 100000
```

```python
# Friendly testing area. 
# This cell is NOT graded. 
from random import randint

print("Input  Base  Converted")
print("=====  ====  =========")
for i in range(20): 
    n = randint(0,max)      # Define max in the previous cell 
    d = convert(n,base)     # Define base in the previous cell 
    print("{:<7}{:<6}{:<11}".format(n,base,d))
```

Output:
```
Input  Base  Converted
=====  ====  =========
31075  28    1bhn       
81003  28    3j8r       
46061  28    22l1       
52015  28    2a9j       
83373  28    3m9h       
81915  28    3kdf       
4761   28    621        
36632  28    1ik8       
94530  28    48g2       
76005  28    3cqd       
33916  28    1f78       
30331  28    1aj7       
23385  28    11n5       
66791  28    315b       
63200  28    2oh4       
84321  28    3nfd       
21081  28    qop        
69185  28    346p       
25340  28    1490       
99720  28    4f5c
```

<!-- grade_id: cell-b3d10603696d09f4, points: 1, grade: true, solution: false -->
```python
# In each of the tests (including 4 hidden tests) convert(number,base) is such that 
# number is a non-negative integer and base in an integer in the interval [2,32]
assert convert(0,7) == '0'
assert convert(17,8) == '21'
assert convert(854,2) == '1101010110'
assert convert(183,5) == '1213'
assert convert(388,11) == '323'
assert convert(3496,32) == '3d8'
assert convert(5961,28) == '7gp'
assert convert(78932148,25) == '821gan'
```

**(b) (1 mark)** Write a function `revert` that, given pair `(num_string,base)`, where `base` is a number in the set  $\{2,\dots,32\}$  and `num_string` is  a string of digits in the set $\{0,\dots,m-1\}$, returns the integer (in decimal) of which `num_string` is the `base`-ary representation. (So, for example, `revert('4b',12)` should return `59`.) 

<!-- grade_id: cell-310ec161adb2380d, points: 0, grade: false, solution: true -->
```python
def revert(num_string,base): 
    # YOUR CODE HERE
    #raise NotImplementedError()
    wt_message = 'Input of wrong type'
    wt_base_message = 'Base imput of wrong type'
    small_base_message = 'Base imput is less than 2'
    large_base_message = 'Base imput is bigger than 32'
    if not type(num_string) is str: 
        print(wt_message)
        return wt_message
    if not type(base) is int:
        print(wt_base_message)
        return wt_base_message
    if base < 2:
        print(small_base_message)
        return small_base_message
    if base > 32:
        print(large_base_message)
        return large_base_message
    #Now for the main part
    
    size = len(num_string)
    product = 1                
    current_n = 0
    
    for i in range(size):
        bit = (num_string[size - i - 1]) 
        if not str. isdigit(bit):
            bit = ord(bit) - 87
        current_n += int(bit) * product
        product *= base
            
    return current_n


```

```python
# Change the value of base and test your functions with this value
# in the next cell 
base = 26
```

```python
# Some friendly checks 
# This cell is NOT graded. 
from random import randint
print("Input  Base  Converted  Reverted")
print("=====  ====  =========  ========")
for i in range(20): 
    n = randint(0,5000)
    d = convert(n,base)     # Define base in the previous cell
    k = revert(d,base)
    print("{:<7}{:<6}{:<11}{:<8}".format(n,base,d,k))
```

Output:
```
Input  Base  Converted  Reverted
=====  ====  =========  ========
3827   26    5h5        3827    
2202   26    36i        2202    
1183   26    1jd        1183    
4816   26    736        4816    
3534   26    55o        3534    
1529   26    26l        1529    
4751   26    70j        4751    
1308   26    1o8        1308    
4475   26    6g3        4475    
432    26    gg         432     
1479   26    24n        1479    
4810   26    730        4810    
4672   26    6ni        4672    
1832   26    2ic        1832    
137    26    57         137     
2293   26    3a5        2293    
2706   26    402        2706    
308    26    bm         308     
3044   26    4d2        3044    
4859   26    74n        4859
```

<!-- grade_id: cell-f95542edb2dc96de, points: 1, grade: true, solution: false -->
```python
# In each of the tests (including 4 hidden tests) revert(num_string,base) is such that 
# num_string is a valid base-ary string representation. For example below '1213' is a 
# valid 5-ary string representation (of the number 854 in decimal). 
# Also, (in each test) base in an integer in the interval [2,32]
assert revert('0',5) == 0
assert revert('1101010110',2) == 854
assert revert('1213',5) == 183
assert revert('323',11) == 388
assert revert('3d8',32) == 3496
assert revert('7gp',28) == 5961
assert revert('e2d89c5',17) == 341893567
```

 # Question 2: Perfect numbers 
 
It is an open problem as to whether any odd perfect numbers exist. On the other hand we know that, for any even number $n$, $n$ is perfect if and only if $n = (2^p - 1) \cdot 2^{p-1}$
and $2^p - 1$ is prime. Note that the if $2^p - 1$ is prime then $p$ is also prime. (Remember also that a prime of the form $2^p - 1$ is called a *Mersenne prime*.) 

**(1 mark)** Write down a function `list_even_perfect` that, given any non-negative integer `upper` as input, returns a list of all the even perfect numbers in the interval `[0,upper]`. 

Note that you will need the `is_prime` function that we defined in the lectures and hence this is imported from the `miller_rabin` module. (So make sure that the file `miller_rabin.py` is in the folder that you are presently working in.) 

<!-- grade_id: cell-fbe3548db00692cc, points: 0, grade: false, solution: true -->
```python
from miller_rabin import is_prime
    
def list_even_perfect(upper): 
    # YOUR CODE HERE
    #raise NotImplementedError()
    if upper == 0:
        return []
    else:
        primes = []
        for x in range(upper+1):
            B = is_prime(x)
            if B == True:
                primes.append(x)
        perf_num_list = [] 
        for p in primes:
            q = ((2**p)-1)*(2**(p-1))
            i = (2**p)-1
            j = is_prime(i)
            if j == True and q <= upper:
                perf_num_list.append(q)
    return perf_num_list

print(list_even_perfect(6))
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-14-15d92866f87c> in <module>
----> 1 from miller_rabin import is_prime
      2 
      3 def list_even_perfect(upper):
      4     # YOUR CODE HERE
      5     #raise NotImplementedError()

ModuleNotFoundError: No module named 'miller_rabin'
```

<!-- grade_id: cell-76217b8142efd508, points: 1, grade: true, solution: false -->
```python
# In each of the tests (including one hidden test) the input is a non-negative integer. 
assert list_even_perfect(0) == []
assert list_even_perfect(6) == [6]
assert list_even_perfect(28) == [6,28]
assert list_even_perfect(85) == [6,28]
assert list_even_perfect(496) == [6,28,496]
assert list_even_perfect(8127) == [6,28,496]
perf_list = [6,28,496,8128,33550336,8589869056,137438691328,
             2305843008139952128,2658455991569831744654692615953842176]
assert list_even_perfect(10**50) == perf_list
```

# Question 3: Primitive Roots

This question follows on from Question 3 of the week 9 tutorial. The functions defined during the tutorial `is_sg_prime` and `is_prim_root_safe` are imported below from the module `primitive_roots`. The function `is_prime` is also imported from the module `miller_rabin`. 
(So make sure that the files `primitive_roots.py` and `miller_rabin.py` are in the folder that you are presently working in.) 

**Reminder.** A prime number $q$ is called a *Sophie Germain prime* if the number
$p := 2 \cdot q + 1$ is also prime. A prime $p$ of this type (i.e. such that 
$p := 2 \cdot q + 1$ where $q$ is also prime) is called a *safe prime*. 

**(a) (1 mark)** Write down a function `find_prim_root_safe` that, given a safe prime `p` as input returns the least primitive root modulo `p`.  If the input is a non-negative integer that is not a safe prime your function should both print and return the string `'Input not a safe prime'`. (Note that if `a` is a primitive root modulo `p`, then `a` $\ge$ `2` by definition.) 

```python
from miller_rabin import is_prime
from primitive_roots import is_sg_prime, is_prim_root_safe
```

<!-- grade_id: cell-d8aebcc14b9bccc0, points: 0, grade: false, solution: true -->
```python

def find_prim_root_safe(p): 
    # YOUR CODE HERE
    #raise NotImplementedError()
    nsp_message = ('Input not a safe prime')
    
    q = p - 1
    qq = q/2
    if q%2 == 0 and is_sg_prime(int(qq)) == True:
        b = 2
        while b in range(2,p):
            if is_prim_root_safe(b,p) == True:
                return b
            else:
                b += 1
                
    else:
        return nsp_message

print(find_prim_root_safe(479))
```

<!-- grade_id: cell-971c0548fdfca7be, points: 1, grade: true, solution: false -->
```python
# In each of the tests (including 5 hidden tests) the input is a non-negative integer. 
assert find_prim_root_safe(0) == 'Input not a safe prime'
assert find_prim_root_safe(479) == 13
assert find_prim_root_safe(719) == 11
assert find_prim_root_safe(119759) == 17
assert find_prim_root_safe(2000000579) == 2
assert find_prim_root_safe(2000003519) == 13
assert find_prim_root_safe(2000006543) == 5
assert find_prim_root_safe(31) == 'Input not a safe prime'
assert find_prim_root_safe(1000004123) == 'Input not a safe prime'
```

**(b) (1 mark)** Write down a function `enum_safe_primes` which, given input `(lower,n)`, returns a dictionary containing `n` many  `key:value` pairs where the keys are the first `n` Sophie Germain primes of size at least `lower` and the `value` associated with each `key`  is the pair `(p,a)`,   where `p` is the safe prime corresponding to the `key`, and `a` is the least primitive root modulo `p`. (Note that the round brackets of `(p,a)` indicate that this is a pair which is, by definition,  a $2$-tuple and NOT a list.)

<!-- grade_id: cell-7675ba26501c875e, points: 0, grade: false, solution: true -->
```python
def enum_safe_primes(lower,n): 
    # YOUR CODE HERE
    #raise NotImplementedError()
    key_no = []
    for i in range(lower,n*lower):
        if is_sg_prime(i) == True:
            if len(key_no) < n:
                key_no.append(i)                    
    p = []
    a = []
    d = {}
    for n in range(lower,n+1):
        p[n] = key_no[n]*2 + 1 
        a[n] = find_prim_root_safe(p[n])

        d[key_no[n]] = (p[n],a[n])
        
    return d    

print(enum_safe_primes(350,4))
```

<!-- grade_id: cell-4affc864599486f2, points: 1, grade: true, solution: false -->
```python
# In each of the tests (including 2 hidden tests) the inputs are non-negative integers.
assert enum_safe_primes(0,0) == {}
assert enum_safe_primes(0,1) == {2: (5, 2)}
assert enum_safe_primes(1,1) == {2: (5, 2)}
assert enum_safe_primes(1,5) == {2: (5, 2), 3: (7, 3), 5: (11, 2), 
                                 11: (23, 5), 23: (47, 5)}
assert enum_safe_primes(350,4) == {359: (719, 11), 419: (839, 11), 
                                   431: (863, 5), 443: (887, 5)}
assert enum_safe_primes(3**15,0) == {}
assert enum_safe_primes(3**15,2) == {14348963: (28697927, 5), 14349059: (28698119, 7)}
```

# Question 4: message coding and decoding 

The purpose of this exercise is to define coding and decoding functions in precisely the same way as in the Preliminaries of Lecture 18 BUT using decimal representation instead of binary representation. You will work with characters whose ASCII number is in the range $0$ to $255$ (generated by `range(256)`). In the lecture notes we work with $8$-bit binary representations of the ASCII number of a character. Here you will work with the $3$-digit decimal representation of the ASCII number using the zero character `'0'` to pad out $1$-digit or $2$-digit numbers. (Accordingly `'009'`,`'037'`, `'075'` and `'248'` are the decimal string representations of the characters `'\t'`,`'%'`, `'K'` and `'ø'` respectively.) 

**(a) (1 mark)** Define a function a function `char_to_dec` that, given character `c` (whose ASCII number is in the specified range) as input, outputs the $3$-digit decimal (string) representation of `c`. 

<!-- grade_id: cell-b597f0c6e9f42ec3, points: 0, grade: false, solution: true -->
```python
def char_to_dec(c): 
    # YOUR CODE HERE
    #raise NotImplementedError()
    n = ord(c)
    n_string = ''
    n_string = n_string + str(n)
    if len(n_string) == 2:
        n_string = '0' + n_string
    if len(n_string) == 1:
        n_string = '00' + n_string
        
    return n_string
#print(char_to_dec('h'))

print(char_to_dec)
```

```python
# Friendly testing area. 
# This cell is NOT graded

print("ASCII    Character    Check")
print("=====    =========    =====")
for n in range(32,256):
    c = chr(n)
    c_dec = char_to_dec(c)
    print("{:<9d}{:13s}{:5s}".format(n,c,c_dec))
```

<!-- grade_id: cell-d7e55dc845ec703d, points: 1, grade: true, solution: false -->
```python
# In each of the tests (including 4 hidden tests) the inputs are characters whose 
# ASCII number is in the interval [32,255]
assert char_to_dec(' ') == '032'
assert char_to_dec('#') == '035'
assert char_to_dec('@') == '064'
assert char_to_dec('K') == '075'
assert char_to_dec('}') == '125'
assert char_to_dec('Ä') == '196'
assert char_to_dec('à') == '224'
assert char_to_dec('ý') == '253'
```

**(b) (1 mark)** Write down a function `dec_convert_to_integer` that behaves like `convert_to_integer` from the Preliminaries of  Lecture 18 but uses `char_to_dec` instead of `char_to_byte`. I.e. `dec_convert_to_integer` takes as input a string `a_text` and outputs a (unique) integer using the decimal method from above (instead of the binary method from the lecture).  

<!-- grade_id: cell-4154c3c598e15089, points: 0, grade: false, solution: true -->
```python
def dec_convert_to_integer(a_text): 
    # YOUR CODE HERE
    #raise NotImplementedError()
    n_string = '1'
    for letter in a_text:
        n_string = n_string + char_to_dec(letter)
    print("The decimal representation of this message is:") 
    print(n_string)
    return int(n_string)

#print(dec_convert_to_integer('abc'))
```

<!-- grade_id: cell-cb9d73b9e7d28ebf, points: 1, grade: true, solution: false -->
```python
# In each of the tests (including 2 hidden tests) the input is a string comprising 
# a message written in a European language.(So  all characters used have ASCII number
# in the interval [32,255])
message_one = "This is a message that will be coded and decoded."
message_two = "Nous avons réussi à communiquer en secret.\n"
message_two += "Mais faites attention: il y a des espions dans les environs."
print("Message one")
print("===========")
print(message_one,"\n")
print("Message two (in French)")
print("=======================")
print(message_two,"\n")
print("Now testing whether your function encodes these ") 
print("messages correctly...\n\n",)
# The integers corresponding to each of the above messages
dec_code_one = int('''
1084104105115032105115032097032109101115115
0971031010321161040971160321191051081080320
9810103209911110010110003209711010003210010
1099111100101100046
'''.replace('\n',''))
dec_code_two = int('''
1078111117115032097118111110115032114233117
1151151050322240320991111091091171101051131
1710111403210111003211510109911410111604601
0077097105115032102097105116101115032097116
1161011101161051111100580321051080321210320
9703210010111503210111511210511111011503210
0097110115032108101115032101110118105114111
110115046
'''.replace('\n',''))
# Testing your output
assert dec_convert_to_integer(message_one) == dec_code_one
assert dec_convert_to_integer(message_two) == dec_code_two
```

**(c) (1 mark)** Write down a function `dec_convert_to_text` which takes as input a positive integer and returns the string (i.e. a piece of text)  that this number is a code for under the method implemented by your function `dec_convert_to_integer`. 

<!-- grade_id: cell-41daa9cdbed9158a, points: 0, grade: false, solution: true -->
```python
def dec_convert_to_text(n):
    # YOUR CODE HERE
    #raise NotImplementedError()
    wt_message = 'Input of wrong type'
    ng_message = 'Input is negative'
    if not type(n) is int: 
        print(wt_message)
        return wt_message
    if n < 0: 
        print(ng_message)
        return ng_message
    
    
    n_string = str(n)[1:] 
    text = ''
    length = len(n_string)
    for i in range(0,length,3):
        three_string = n_string[i:i+3]
        text += chr(int(three_string))
    return text    

#print(dec_convert_to_text(1084104105115032))
```

<!-- grade_id: cell-a1fa6771bb5fbe9c, points: 1, grade: true, solution: false -->
```python
# In each of the tests (including 2 hidden tests) the input is a 
# positive integer. Note that we use message_one and message_two 
# from the testing area of part (b) above. So make sure that the 
# testing  area of part (b) is executed first - otherwise message_one 
# and message_two will not be defined. 
# We start  by printing out these messages. 
print("Message one after coding and successful decoding should be...")
print("===========================================================")
print(message_one,"\n")
print("Message two after coding and successful decoding should be...")
print("===========================================================")
print(message_two,"\n")
print("Now testing whether your function does the decoding part") 
print("of this correctly...\n\n",)
# Now we test that your function does this for (i.e. decodes correctly) the 
# integers dec_code_one and dec_code_two from the testing area of part (b)
assert dec_convert_to_text(dec_code_one) == message_one
assert dec_convert_to_text(dec_code_two) == message_two
```

# Question 5: encoding pairs of integers

In the study of the algorithmic content of mathematics it is essential to have tools for encoding finite objects/structures, such as finite sequences, as integers. In this exercise we will study a pairing function - i.e. a function which encodes 2-element sequences. 

**Notation.** We use $\mathbb{N}$ to denote the set of non-negative integers 
$\{n \;|\;n \in \mathbb{Z} \;\&\; n \ge 0 \}$. 

**(a)** We define the paring function 
$f : \mathbb{N} \times \mathbb{N} \rightarrow \mathbb{N}$ by 
$$
f : (x,y) \mapsto \frac{1}{2} \left(x + y\right)\left(x + y + 1\right) + x \,. 
$$

**Note 1.** $f$ is bijective. 

**Note 2.** The formula for $f$ is derived using the infinite grid representation of 
$\mathbb{N} \times \mathbb{N}$ where row $x$  contains all pairs with first component $x$ and column $y$ contains all pairs with second component $y$. In the  diagram below you can see the $9$ pairs in the  bottom left hand corner of this grid. Starting with 
$(0,0)$ we follow the arrows. In this way we take each diagonal and walk down it from the top. The number of pairs on each diagonal is one more than the sum of the entries of each pair on the diagonal. This means that, for any pair $(x,y)$, the number of pairs on previous diagonals is $1 + 2 + \cdots + (x + y)$. Also the number of pairs above it on the same diagonal is $x$. 
Thus we can simply define $f$ on input $(x,y)$ as the function that counts the number of pairs preceding $(x,y)$ in the grid according to this pattern.  In this way we obtain the formula for $f$ shown above. (Hence under $f$ we have that  $(0,0) \mapsto 0$, $(0,1) \mapsto 1$, $(1,0) \mapsto 2, \dots ,(3,0) \mapsto 9$ etc.)  

<img src="pairing_function.jpg" alt="pairs_diagram" style="width: 400px">

(Taken from the book *Computability Theory* by Rebecca Weber, AMS, 2016.) 

**(a) (1 mark)** Write down a function `encode_pair` that implements $f$, i.e. which, given input `a_pair` with `a_pair` = `(x,y)` for some non-negative integers `x`, `y`, returns  the integer 
given by $f\left((x,y)\right)$. 


<!-- grade_id: cell-79b1b7bb5da215dd, points: 0, grade: false, solution: true -->
```python
def encode_pair(a_pair):
    # YOUR CODE HERE
    #raise NotImplementedError()
    for i in a_pair:
        a = a_pair[0]
        b = a_pair[1]
        A = (1/2)*(a+b)*(a+b+1)+a
    return int(A)
#print(encode_pair((6,3)))
```

```python
# Friendly testing  area 
# This cell is NOT graded
pairs = [(0,0),(0,1),(1,0),(0,2),(1,1)]
pairs += [(2,0),(0,3),(1,2),(2,1),(3,0)]
for p in pairs: 
    print("The code of {} is {}".format(p,encode_pair(p)))
```

<!-- grade_id: cell-30a734226012b766, points: 1, grade: true, solution: false -->
```python
# In each of the tests (including 8 hidden tests) the input is 
# a pair of non-negative integers.
assert encode_pair((0,0)) == 0
assert encode_pair((2,1)) == 8
assert encode_pair((6,3)) == 51
assert encode_pair((6,4)) == 61
assert encode_pair((2,40)) == 905
assert encode_pair((123,712)) == 349153
assert encode_pair((65304178,8234117)) == 2703940517826838
```

**(b) (1 mark)** As noted above $f$ is bijective so that $f^{-1}$ is defined. Write down a function 
`decode_to_pair` which, given non-negative integer `n` as input returns the pair given by $f^{-1}(n)$.  

<!-- grade_id: cell-3d59ca6c735a127c, points: 0, grade: false, solution: true -->
```python
def decode_to_pair(n):
    # YOUR CODE HERE
    #raise NotImplementedError()
    pair1 = (0,0)
    if encode_pair(pair1) == n:
        return pair1
    else:
        for i in range(n):
            for j in range(n):
                if encode_pair((i,j)) == n:
                    return (i,j)
        #return pair          
                
                
    
    
#print(decode_to_pair(905))        
```

```python
# Friendly testing  area 
# This cell is NOT graded
for n in range(10): 
    print("{} decoded is the pair {}".format(n,decode_to_pair(n)))
```

<!-- grade_id: cell-95eab1af911a7436, points: 1, grade: true, solution: false -->
```python
# In each of the tests (including 6 hidden tests) the input is a non-negative integer.
# (Also there is no input greater than 20000, as the function code_to_pair may use up a 
# large amount of time and space.)
assert decode_to_pair(0) == (0,0)
assert decode_to_pair(8) == (2,1)
assert decode_to_pair(51) == (6,3)
assert decode_to_pair(61) == (6,4)
assert decode_to_pair(905) == (2,40)
assert decode_to_pair(2135) == (55,9)
assert decode_to_pair(17824) == (58,130)
```

```python

```